# Assignment 2

For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

```python
import pandas as pd
import numpy as np
```

```python
dataset= pd.read_csv('assets/NISPUF17.csv',index_col = 0)
dataset


```

Output:
```
SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  YEAR  \
1       128521     12852     2         NaN  235.916956     1031  2017   
2        10741      1074     2         NaN  957.353840     1068  2017   
3       220011     22001     2         NaN  189.611299     1050  2017   
4        86131      8613     1  675.430817  333.447418     1040  2017   
5       227141     22714     1  482.617748  278.768063     1008  2017   
...        ...       ...   ...         ...         ...      ...   ...   
28461    19141      1914     2         NaN   21.188088     2062  2017   
28462   121401     12140     2         NaN   26.503010     2004  2017   
28463   128831     12883     2         NaN   27.520652     2004  2017   
28464    82921      8292     2         NaN   22.390587     2062  2017   
28465   244621     24462     2         NaN   27.520652     2004  2017   

       AGECPOXR  HAD_CPOX  AGEGRP  ...  XVRCTY2  XVRCTY3  XVRCTY4  XVRCTY5  \
1           NaN         2       1  ...              
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    dict = {}  # initializaiton of dictionary
    total = len(dataset) # total number of children for all category
    d1 = dataset[dataset['EDUC1']== 1] # < 12 years of study
    d2 = dataset[dataset['EDUC1']== 2] # 12 years of study
    d3 = dataset[dataset['EDUC1']== 3] # > 12 years of study but not college graduate
    d4 = dataset[dataset['EDUC1']== 4] # college graduate
    dict["less than high school"] = len(d1)/total # append on dictionery
    dict["high school"] = len(d2)/total
    dict["more than high school but not college"]=len(d3)/total
    dict["college"] = len(d4)/total
    return dict
    # your code goes here
    # YOUR CODE HERE
   # raise NotImplementedError()

```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

```python

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    data_flu_breast = dataset[['CBF_01','P_NUMFLU']]
#print(dat1.CBF_01.unique())
#print(dat1.EDUC1.unique())

    data_flue_breast = data_flu_breast.dropna()
    dat_know = data_flue_breast[data_flue_breast['CBF_01'] == 1]
    dat_do_not_know = data_flue_breast[data_flue_breast['CBF_01'] == 2]
    return (dat_know['P_NUMFLU'].sum()/len(dat_know),dat_do_not_know['P_NUMFLU'].sum()/len(dat_do_not_know))
    # YOUR CODE HERE
    #raise NotImplementedError()
```

```python
average_influenza_doses()
```

Output:
```
(1.8799187420058687, 1.5963945918878317)
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    ratio = {}
    dat3 = dataset[['HAD_CPOX','P_NUMVRC','SEX']].dropna()
    dat_chick= dat3[dat3['HAD_CPOX']==1]
    dat_no_chick = dat3[dat3['HAD_CPOX']==2]
    dat_chick_vac = dat_chick[dat_chick['P_NUMVRC'] != 0]
    dat_no_chick_vac = dat_no_chick[dat_no_chick['P_NUMVRC'] != 0]
    ratio['male'] = len(dat_chick_vac[dat_chick_vac['SEX'] ==1])/len(dat_no_chick_vac[dat_no_chick_vac['SEX']==1])
    ratio['female'] = len(dat_chick_vac[dat_chick_vac['SEX'] ==2])/len(dat_no_chick_vac[dat_no_chick_vac['SEX']==2])
    return ratio
                                         
                                                 
                                         
    # YOUR CODE HERE
    #raise NotImplementedError()
```

```python
chickenpox_by_sex()
```

Output:
```
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

```python
dat4 = dataset[['HAD_CPOX','P_NUMVRC']].dropna()
dat_chick= dat4[dat4['HAD_CPOX'] != 77]
dat_vac = dat4['P_NUMVRC']
dat_vac = dat4[dat4['HAD_CPOX'] != 77]
dat_vac = dat_vac['P_NUMVRC']
len(dat_vac)

```

Output:
```
15286
```

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    #df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
                  # "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})

    # here is some stub code to actually run the correlation
    #corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    
    # just return the correlation
    #return corr

    # YOUR CODE HERE
    dat4 = dataset[['HAD_CPOX','P_NUMVRC']].dropna()
    dat_chick1= dat4[dat4['HAD_CPOX'] != 77]
    dat_chick = dat_chick1['HAD_CPOX']
    
    #dat_vac = dat4[dat4['HAD_CPOX'] != 77]
    dat_vac = dat_chick1['P_NUMVRC']
    corr,pval = stats.pearsonr(dat_chick,dat_vac)
    #raise NotImplementedError()
    return corr
```

```python
corr_chickenpox()
```

Output:
```
0.07044873460147986
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

```python

```